# Assignment 4 - Bagging and Hyperparameter tuning  

<img src="./images/bags.jpg" alt="The bagging operation" width="400" align='center'/>

## The Spam vs. not-Spam Dataset
<img src="./images/spam_or_not_spam.png" alt="Spam vs. not-Spam" width="300" align='center'/>

### About the (Spam vs. not-Spam) dataset:
* The dataset tags email messages as spam or not-spam.
* Classes (=categories): there are two possible classes: not_spam, spam
* Attributes (=features): there are 57 features, including two types of features:
  * word frequencies - the feature name contains a word, with the suffix '_wordFreq' e.g.: 'make_wordFreq'  
    * The word could appear 0 times in a specific email, once or a few times.
    * The values of this feature could be between 0 and 1 (the frequency is relative).
  * Capital Letter pattern attributes - 3 features regarding capital letter patterns. e.g.: 'capitalLet_long' - the length of the longest sequence of capital letters in the email.

### The assignment has to do with the (Spam vs. not-Spam) dataset
* This assignment is a bigger assignment and is graded accordingly (20 points total)
  * It includes different methods to implement, related to the Bagging algorithm 
* At the end of the assignment, there is an optional bonus method to implement (additional 5 points)
  * It also includes an , regarding hyperparameters tunig 
* Note: the maximum possible points for all the assignments together are - 70 points 
  * The extra 30 points come from the test
* Note: previous assignments and class exercises, could assist you with similar materials

## General instructions:
Please Read the following instructions, and make sure you follow all steps, before submiting your assignment via moodle:
- Put your code one line after the '# YOUR CODE HERE' remark
- You must remove the 'raise NotImplementedError()' exception raising line (which appears a line after the above remark).
    * if you do not remove this line, it will be a sign that you didn't implement that code, and we won't be able to check your work
- Do NOT remove any other line in this notebook    
- you also need to implement the two functions 'myName', 'myId'
    * myName - you need to return your full name as a string
    * myId - you need to return your ID number as a an integer    
- When you want to check your work, select the 'Cell' --> 'Run All' menu
    * before performing your filnal test, we suggest to clear previous output (by selecting 'Cell' --> 'All output' --> 'Clear' menu) and then reperform the final execution ('Cell' --> 'Run All') of the code.
    * after performing 'Run All', make sure there are no exceptions thrown and that the output is as you expected.
- Don't forget to save your work, by clicking the 'save' icon (in the upper left part of the screen), or by selecting the 'File' --> 'Save and checkpoint' menu item. 
- Do NOT change the file name of this notebook
- The work is done individualy, for each student
- Each student needs to submit his/her assignment through moodle.
- Submit the python notebook only (not any additional possible files) 
<br/><br/>
Good Luck :-)<br/>
The courses staff

### Part 1 - Student information methods:
The following 2 cells consist of 2 methods which aim to validate your details.<br/>
Please <b>make sure to implement <u>both</u> of them</b>

<!-- grade_id: myName-Method, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# method name: 
# --------------------------------------------------------

# --------------------------------------------------------
# The following is expected:
# --- the method needs to return your full name (as a string)
# example: assume your name is John Smith:
# return 'John Smith'
# ------------
# return value:
# - your full name (as a string)
# ---------------------
# --- General notes:
# Don't forget to remove the 'raise NotImplementedError()' line.
# In the line following the '# YOUR CODE HERE' line, add your code
# ---------------------
def myName():
    # ADD the return statement in the LINE AFTER: '# YOUR CODE HERE'
    # REMOVE THE LINE: raise NotImplementedError()
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: myIdMethod, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# method name: 
# --------------------------------------------------------

# --------------------------------------------------------
# The following is expected:
# --- the method needs to return your ID number (as an integer number)
# example: assume your ID number is 1234:
# return 1234 
# ------------
# return value:
# - your ID number (as an integer number)
# ---------------------
# --- General notes:
# Don't forget to remove the 'raise NotImplementedError()' line.
# In the line following the '# YOUR CODE HERE' line, add your code
# ---------------------
def myId():
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# --------------------------------------------------------
# TEST ONLY CELL
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# Important Notes:
# - This test cell MUST NOT raise any exception
# - It must pass all tests 
# - Do NOT change or delete this cell
# --------------------------------------------------------

# --------------------------------------------------------
# Test of:
## 'myName' and 'myId' implementation methods
# It tests the correctness their implementation
# Important Note: additional test might also be taken from our side.
# --------------------------------------------------------
aName = myName()
aId = myId()
assert aName is not None, 'no return student name'
assert type(aName) is str, "name is not a string: %r" % aName
print ("'myName' method validation is successfull!")
print ("your name is: %s" %(aName))
print ("--------------------------------------")
assert aId is not None, 'no return student id'
assert type(aId) is int, "id is not an integer: %r" % aId
print ("'myId' method validation is successfull!")
print ("your id is: %d" %(aId))
## END OF 'myName' and 'myId' implementation validation
# --------------------------------------------------------
```

### Part 2 - Loading dataset - maximum possible points: 3
The following cells perform 5 things:
* Preceding step - import packages
* step 1 - load the 'spam vs not-spam' dataset ----> Student's implementation - total 1 point
* step 2 - split dataframe to X (feature vectors) and y (classes) ----> Student's implementation - total 1 point 
* step 3 - change classes values from string to number ----> Student's implementation - total 1 point 
* step 4 - split the dataset to a train-set and a test-set ----> run only

### Preceding Step - import packages
This step is necessary in order to use external packages. <br/><br/>
<u>Some of the exteranl packages include</u>: 
* pandas - which we use mainly for dataframes and series
* numpy - which we use for advance operations, such as unique values of arrays
* random - which we could use to select random value (like we studied in earlier exercises)

```python
# --------------------------------------------------------
# IMPORT (PACKAGES) CELL
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# - Do NOT change or delete the following imports:
import sys
import os
import pathlib 
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
import random
get_ipython().run_line_magic('matplotlib', 'inline')
# --------------------------------------------------------

# --------------------------------------------------------
# Note: do NOT add imports, such as sklearn, which could
# answer the following questions.
# --- add your imports here IF NEEDED:


# --------------------------------------------------------
```

#### Part 2 - step 1 - load the 'spam vs not-spam' dataset ----> Student's implementation - total 1 point
<img src="./images/load_dataframe.jpg" alt="load dataframe" width="100" align='center'/>

```python
# ------------------------------------------
# Add assistance code here IF NEEDED:



# ------------------------------------------
```

<!-- grade_id: loading_dataframe, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 1
# --------------------------------------------------------

# --------------------------------------------------------
# method name: loadDataset 
# --------------------------------------------------------

# --------------------------------------------------------
# The following is expected:
# --- the method needs to return the dataset from the input csv file.
# ------------
# input parameters:
# - 'fileName' - full path of the csv fileName 
# ------------
# return value:
# - a dataframe structure containing the dataset
# ---------------------
# --- General notes:
# Don't forget to remove the 'raise NotImplementedError()' line.
# In the line following the '# YOUR CODE HERE' line, add your code
# ---------------------
# --- Additional notes:
# - You must use the parameters you've got. 
# - Do NOT define the parameters in the method.
# --------------------- 
def loadDataset(fileName):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# --------------------------------------------------------
# - add your tests here IF NEEDED:
# Note: any additional tests are for your use only.
# Note: We depend only on tests we wrote.
# ------


# --------------------------------------------------------
```

<!-- grade_id: loading_dataframe_test, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# Important Notes:
# - This test cell MUST NOT raise any exception
# - It must pass all tests to get the point ----- #
# - Do NOT change or delete this cell
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 1
# --------------------------------------------------------

# --------------------------------------------------------
# Test of:
### --- Graded tests for the 'loadDataset' method 
# Important Note: additional test might also be taken from our side.
# --------------------------------------------------------

# -----------------------------------------------------------------
folderName = '.' + os.sep + 'data' 
datasetCsvFileName = folderName + os.sep + 'spam_notSpam.csv'
# --------------------- 
print ("Testing your implementation of the 'loadDataset' method ...")
dataset_forTesting = loadDataset(datasetCsvFileName)
assert dataset_forTesting is not None, 'spam_notSpam object not initialized'
assert isinstance(dataset_forTesting, pd.DataFrame), 'spam_notSpam object is not a dataframe'
assert dataset_forTesting.empty == False, 'spam_notSpam dataframe object is empty'
assert len(dataset_forTesting.columns)==58, 'spam_notSpam dataframe object is missing columns'
assert len(dataset_forTesting.index)==4601, 'spam_notSpam dataframe object is missing rows'
print ("----> The 'spam vs not-spam' dataframe object was loaded successfuly :-) \n")
dataset_forTesting=None
print ('The beginning (the head) of the dataframe:')
print ('\nNote: Additional Tests might be executed on our side')
loadDataset(datasetCsvFileName).head()
# --------------------------------------------------------
```

#### Part 2 - step 2 - split dataframe to X (feature vectors) and y (classes) ----> Student's implementation - total 1 point

```python
# ------------------------------------------
# Add assistance code here IF NEEDED:



# ------------------------------------------
```

<!-- grade_id: seperate_to_x_and_y, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 1
# --------------------------------------------------------

# --------------------------------------------------------
# method name: separateTo_X_and_y 
# --------------------------------------------------------

# --------------------------------------------------------
# The following is expected:
# ---- The method should separate the input dataset into two parts:
# 1. X (feature vectors)  
# 2. y (categories) 
# ------------
# input parameters:
# - dataset -  a dataframe structure, containing the dataset.
# - classColName - the column name (string) containing the categories
# ------------
# return values (comma separated):
# - X_featureVectores - a dataframe containing all feature vectors. 
#                       It should contain the input dataframe after removing the class column.
#                       The index of the X_featureVectores should be the same as the input dataframe parameter.
# - y_Categories      - a series of containing all class values per instance.
#                       The index of the y_Categories series should be the same as the input dataframe.
# ---------------------
# --- General notes:
# Don't forget to remove the 'raise NotImplementedError()' line.
# In the line following the '# YOUR CODE HERE' line, add your code
# ---------------------
# --- Additional notes:
# - You must use the parameters you've got. 
# - Do NOT define the parameters in the method.
# --------------------- 
def separateTo_X_and_y(dataset, classColName):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# --------------------------------------------------------
# - add your tests here IF NEEDED:
# Note: any additional tests are for your use only.
# Note: We depend only on tests we wrote.
# ------


# --------------------------------------------------------
```

<!-- grade_id: separateTo_X_and_y-test, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# Important Notes:
# - This test cell MUST NOT raise any exception
# - It must pass all tests to get the point ----- #
# - Do NOT change or delete this cell
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 1
# --------------------------------------------------------

# --------------------------------------------------------
# Test of:
### --- Graded tests for the 'separateTo_X_and_y' method 
# Important Note: additional test might also be taken from our side.
# --------------------------------------------------------

# --------------------------------------------------------
folderName = '.' + os.sep + 'data' 
datasetCsvFileName = folderName + os.sep + 'spam_notSpam.csv'
# --------------------------------------------------------

# --------------------------------------------------------
print ("Testing your implementation of the 'separateTo_X_and_y' method ...")
# --------------------- 
dataset_forTesting = loadDataset(datasetCsvFileName)
X_vectors, y_categories = separateTo_X_and_y(dataset_forTesting, 'class')
# ---------------------
sameIndexes = lambda obj1,obj2: False not in (obj1.index.tolist()[n] == obj2.index.tolist()[n] for n in range(len(obj1.index))) 
assert X_vectors is not None, 'X_vectors object not initialized'
assert isinstance(X_vectors, pd.DataFrame), 'X_vectors object is not a dataframe'
assert X_vectors.empty == False, 'X_vectors dataframe object is empty'
assert sameIndexes(X_vectors,dataset_forTesting), 'X_vectors should share the same indexes as the dataset'
assert y_categories is not None, 'y_categories object not initialized'
assert isinstance(y_categories, pd.Series), 'y_categories object is not a series'
assert y_categories.empty == False, 'y_categories dataframe object is empty'
assert sameIndexes(X_vectors,y_categories), 'X_vectors should share the same indexes as y_categories'
print ("----> The 'separateTo_X_and_y' test passed successfully :-) \n")
print ('\nAdditional Tests might be executed on our side')
dataset_forTesting = None
X_vectors = None
y_categories = None
# --------------------------------------------------------

```

#### Part 2 - step 3 - change classes values from string to number ----> Student's implementation - total 1 point

```python
# ------------------------------------------
# Add assistance code here IF NEEDED:



# ------------------------------------------
```

<!-- grade_id: categories-to-nums, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 1
# --------------------------------------------------------

# --------------------------------------------------------
# method name: separateTo_X_and_y 
# --------------------------------------------------------

# --------------------------------------------------------
# The following is expected:
# ---- The method changes the input string categories to int categories.
#      'not_spam' --> 0
#      'spam'     --> 1
# ------------
# input parameters:
# - y_Categories -  a dataframe structure, containing the dataset.
# - positiveClassName - the name of the category (string) of the positive class
#                       * if the value of a cell does not equal positiveClassName, 
#                         it should be considered as negative.
# ------------
# return value: 
# - y_NumCategories - a series of containing all class values with numeric values per instance.
#                       * each cell value, which equals the negativeClassName will recieve a value of 0 in the
#                         output series.
#                       * each cell value, which equals the positiveClassName will recieve a value of 1 in the
#                         output series.
#                       * the index should be the same of the input series.
# ---------------------
# --- General notes:
# Don't forget to remove the 'raise NotImplementedError()' line.
# In the line following the '# YOUR CODE HERE' line, add your code
# ---------------------
# --- Additional notes:
# - You must use the parameters you've got. 
# - Do NOT define the parameters in the method.
# --------------------- 
def datasetCategoriesToNums(y_Categories, positiveClassName):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# --------------------------------------------------------
# - add your tests here IF NEEDED:
# Note: any additional tests are for your use only.
# Note: We depend only on tests we wrote.
# ------


# --------------------------------------------------------
```

<!-- grade_id: categories-to-nums-test, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# Important Notes:
# - This test cell MUST NOT raise any exception
# - It must pass all tests to get the point ----- #
# - Do NOT change or delete this cell
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 1
# --------------------------------------------------------

# --------------------------------------------------------
# Test of:
### --- Graded tests for the 'datasetCategoriesToNums' method 
# Important Note: additional test might also be taken from our side.
# --------------------------------------------------------

# --------------------------------------------------------
folderName = '.' + os.sep + 'data' 
datasetCsvFileName = folderName + os.sep + 'spam_notSpam.csv'
# --------------------------------------------------------

# --------------------------------------------------------
print ("Testing your implementation of the 'datasetCategoriesToNums' method ...")
# --------------------- 
dataset_forTesting = loadDataset(datasetCsvFileName)
X_vectors, y_categoriesB4 = separateTo_X_and_y(dataset_forTesting, 'class')
y_categories = datasetCategoriesToNums(y_categoriesB4, 'spam')
# --------------------- 
allValidVals = lambda arrValues,validVals: False not in (val in validVals for val in arrValues) 
sameIndexes = lambda obj1,obj2: False not in (obj1.index.tolist()[n] == obj2.index.tolist()[n] for n in range(len(obj1.index))) 
assert y_categories is not None, 'y_categories object not initialized'
assert isinstance(y_categories, pd.Series), 'y_categories object is not a series'
assert y_categories.empty == False, 'y_categories dataframe object is empty'
assert sameIndexes(y_categories,y_categoriesB4), 'y_categories should share the same indexes as y_categoriesB4'
assert allValidVals(np.unique(y_categories.values),(0,1)), 'classes should be only 0 or 1 (as an integer number)'
print ("----> The 'datasetCategoriesToNums' test passed successfully :-) \n")
print ('\nAdditional Tests might be executed on our side')
dataset_forTesting = None
X_vectors = None
y_categories = None
# --------------------------------------------------------
```

#### Part 2 - step 4 - split the dataset to a train-set and a test-set ----> run only
<img src="./images/train-test-split.png" alt="train-test split" width="100" align='center'/>

```python
# --------------------------------------------------------
# NO IMPLEMENTAION - DO NO CHANGE
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# method name1: trainTestSplit
# --------------------------------------------------------

# --------------------------------------------------------
# What does the method do?
# --- the method split the input dataset into train and test. 
#     - It does so using the sklearn built in method
# ------------
# input parameters:
# - X_vectors - a dataframe containing all feature vectors. 
# - y_categories - a series of containing all class values per instance.
# - test_size_ratio - a number (0<number<1) of the wanted ratio of the dataset out of the dataset 
# - rand_state - a number, in order to guarantee reproducible results 
# ------------
# return values (comma separated):
# - X_train -  a dataframe containing all feature vectors of the train set
# - X_test -  a dataframe containing all feature vectors of the test set
# - y_train - a series of containing all class values per train instance
# - y_test - a series of containing all class values per test instance
# --------------------- 
def trainTestSplit(X_vectors, y_categories, test_size_ratio, rand_state):
    return train_test_split(X_vectors, y_categories, test_size=test_size_ratio, random_state=rand_state)


# --------------------------------------------------------
# method name2: reAttachTrainSet
# --------------------------------------------------------

# --------------------------------------------------------
# What does the method do?
# --- the method split the input dataset into train and test. 
#     - It does so using the sklearn built in method
# ------------
# input parameters:
# - X_train -  a dataframe containing all feature vectors of the train set
# - y_train - a series of containing all class values per train instance
# ------------
# return value:
# - train_set - a reattached dataframe, consisting both feature vectors and categories.
# --------------------- 
def reAttachTrainSet(X_train, y_train):
    return pd.concat((X_train, y_train), axis=1)

# --------------------------------------------------------
# Temporary imports:
import sklearn
from sklearn.model_selection import train_test_split
# --------
folderName = '.' + os.sep + 'data' 
datasetCsvFileName = folderName + os.sep + 'spam_notSpam.csv'
# --------------------- 
dataset_forTesting = loadDataset(datasetCsvFileName)
X_vectors, y_categories = separateTo_X_and_y(dataset_forTesting, 'class')
y_categories = datasetCategoriesToNums(y_categories, 'spam')
X_train, X_test, y_train, y_test = trainTestSplit(X_vectors, y_categories, 0.2, 25)
# --------------------------------------------------------

# --------------------------------------------------------
if 'datasets' in sys.modules:
    del (datasets)
if 'train_test_split' in sys.modules:
    del (train_test_split)
sys_modules = list(sys.modules.keys())
for mdl in sys_modules:
    if mdl.startswith('sklearn.'):
        del(sys.modules[mdl]) 
del (sklearn)
if 'sklearn' in sys.modules:
    del (sys.modules['sklearn'])
# --------------------------------------------------------

# --------------------------------------------------------
# display train-test split information
# --------------------------------------------------------
print ('Information after train-test split:')
print('The train-set includes %d instances and %d corresponding categories\n' %(X_train.shape[0],y_train.shape[0]))
print('The test-set includes %d instances and %d corresponding categories\n' %(X_test.shape[0],y_test.shape[0]))

# --------------------------------------------------------
## concatinate the X_train and y_train:
# --------------------------------------------------------
train_set = reAttachTrainSet(X_train, y_train)

# --------------------------------------------------------
dataset_forTesting = None
X_vectors = None
y_categories = None
X_train = None
X_test = None
y_train = None
y_test = None
# --------------------------------------------------------

# --------------------------------------------------------
# Display the first few rows of the training-set:
# --------------------------------------------------------
print('First few rows of unified train-set:')
train_set.head()
```

### Part 3 - Building a Bagging Classifier
<img src="./images/bagging-classifier.jpeg" alt="bagging-classifier" width="400" align='center'/>

### Part 3a - Bootstrap - maximum possible points: 8
The following cells perform 2 methods:
* step 1 - instance bootstrap ----> Student's implementation - total 4 point
* step 2 - feature bootstrap ----> Student's implementation - total 4 point

#### Part 3a - step 1 - instance bootstrap ----> Student's implementation - total 4 point

```python
# ------------------------------------------
# Add assistance code here IF NEEDED:



# ------------------------------------------
```

<!-- grade_id: instance-bootstrap, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 4
# --------------------------------------------------------

# --------------------------------------------------------
# method name: instanceBootStrap
# --------------------------------------------------------

# --------------------------------------------------------
# What does the method do?
# --- the method samples the dataset uniformly with replacements. 
#     * this means that there is a chance we could get the same instance more than once
# ------------
# input parameters:
# - X_train - a dataframe containing all training feature vectors. 
# - y_train - a series of containing all training class values per instance.
# - sampleRatio - the ratio of the sampeling out of training set 
#               * we will derive the number of instances to sample from sampleRatio
# ------------
# return values (comma separated):
# - X_trainSampled - a dataframe containing sampled training feature vectors. 
# - y_trainSampled - a series of containing Sampled training class values per instance.
# Note: X_trainSampled should have the same index as y_trainSampled 
# --------------------- 
# --- General notes:
# Don't forget to remove the 'raise NotImplementedError()' line.
# In the line following the '# YOUR CODE HERE' line, add your code
# ---------------------
# --- Additional notes:
# - You must use the parameters you've got. 
# - Do NOT define the parameters in the method.
# --------------------- 
def instanceBootStrap(X_train, y_train, sampleRatio):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# --------------------------------------------------------
# - add your tests here IF NEEDED:
# Note: any additional tests are for your use only.
# Note: We depend only on tests we wrote.
# ------


# --------------------------------------------------------
```

<!-- grade_id: instance-nootStrap-test, points: 4, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# Important Notes:
# - This test cell MUST NOT raise any exception
# - It must pass all tests to get the point ----- #
# - Do NOT change or delete this cell
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 4
# --------------------------------------------------------

# --------------------------------------------------------
# Test of:
### --- Graded tests for the 'instanceBootStrap' method 
# Important Note: additional test might also be taken from our side.

# --------------------------------------------------------
# Temporary imports:
import sklearn
from sklearn.model_selection import train_test_split
# --------
folderName = '.' + os.sep + 'data' 
datasetCsvFileName = folderName + os.sep + 'spam_notSpam.csv'
# --------------------------------------------------------
print ("Testing your implementation of the 'instanceBootStrap' method ...")
dataset_forTesting = loadDataset(datasetCsvFileName)
X_vectors, y_categories = separateTo_X_and_y(dataset_forTesting, 'class')
y_categories = datasetCategoriesToNums(y_categories, 'spam')
X_train, X_test, y_train, y_test = trainTestSplit(X_vectors, y_categories, 0.2, 35)
# --------------------------------------------------------

# --------------------------------------------------------
if 'datasets' in sys.modules:
    del (datasets)
if 'train_test_split' in sys.modules:
    del (train_test_split)
sys_modules = list(sys.modules.keys())
for mdl in sys_modules:
    if mdl.startswith('sklearn.'):
        del(sys.modules[mdl]) 
del (sklearn)
if 'sklearn' in sys.modules:
    del (sys.modules['sklearn'])
# --------------------------------------------------------

# --------------------------------------------------------
allValidVals = lambda arrValues,validVals: False not in (val in validVals for val in arrValues) 
sameIndexes = lambda obj1,obj2: False not in (obj1.index.tolist()[n] == obj2.index.tolist()[n] for n in range(len(obj1.index))) 
X_train_100 = X_train.iloc[:100,:]
y_train_100 = y_train.iloc[:100]
sampleRatio = 0.5
X_trainSampled, y_trainSampled = instanceBootStrap(X_train_100, y_train_100, sampleRatio)
# --------------------------------------------------------

assert X_trainSampled is not None, 'X_trainSampled object not initialized'
assert isinstance(X_trainSampled, pd.DataFrame), 'X_trainSampled object is not a dataframe'
assert X_trainSampled.empty == False, 'X_trainSampled dataframe object is empty'
assert y_trainSampled is not None, 'y_trainSampled object not initialized'
assert isinstance(y_trainSampled, pd.Series), 'y_trainSampled object is not a series'
assert y_trainSampled.empty == False, 'y_trainSampled series object is empty'
assert sameIndexes(X_trainSampled,y_trainSampled), 'X_trainSampled should share the same indexes as y_trainSampled'
assert allValidVals(np.unique(y_categories.values),(0,1)), "y_trainSampled's value should be only 0 or 1 (as an integer number)"
print ("----> The 'instanceBootStrap' test passed successfully :-) \n")
print ('\nAdditional Tests might be executed on our side')
dataset_forTesting = None
X_vectors = None
y_categories = None
X_train = None
X_test = None
y_train = None
y_test = None
X_train_100 = None
y_train_100 = None
sampleRatio = None
X_trainSampled = None
y_trainSampled = None
# --------------------------------------------------------

# --------------------------------------------------------
```

#### Part 3a - step 2 - feature bootstrap ----> Student's implementation - total 4 point

```python
# ------------------------------------------
# Add assistance code here IF NEEDED:



# ------------------------------------------
```

<!-- grade_id: feature-bootstrap, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 4
# --------------------------------------------------------

# --------------------------------------------------------
# method name: featureBootStrap
# --------------------------------------------------------

# --------------------------------------------------------
# What does the method do?
# --- the method samples the featureset uniformly with replacements. 
#     * this means that there is a chance we could choose the same feature twice 
#     * note that in the output dataframe, any feature which was already chosen,
#            will be disregarded. In other words, such a case will result in less
#            output features. For instance if we have 50 input features, and the sample
#            ratio is 0.5, we should expect 25 sampled feature columns. But if one of the features
#            was selected twice, we expect only 24 feature columns.
# ------------
# input parameters:
# - X_instances - a dataframe containing all feature vectors from the dataset.
# - sampleRatio - the ratio of the sampeling out of feature set,
#               * we will derive the number of features from sampleRatio
# ------------
# return value:
# - X_instances_featureSampled - a dataframe containing feature vectors with sampled features. 
#                               * note the instances refer to the same instances as the input dataframe.
#                               * the difference is in the columns (the selected features). 
# --------------------- 
# --- General notes:
# Don't forget to remove the 'raise NotImplementedError()' line.
# In the line following the '# YOUR CODE HERE' line, add your code
# ---------------------
# --- Additional notes:
# - You must use the parameters you've got. 
# - Do NOT define the parameters in the method.
# --------------------- 
def featureBootStrap(X_instances, sampleRatio):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# --------------------------------------------------------
# - add your tests here IF NEEDED:
# Note: any additional tests are for your use only.
# Note: We depend only on tests we wrote.
# ------


# --------------------------------------------------------
```

<!-- grade_id: feature-bootstrap-test, points: 4, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# Important Notes:
# - This test cell MUST NOT raise any exception
# - It must pass all tests to get the point ----- #
# - Do NOT change or delete this cell
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 4
# --------------------------------------------------------

# --------------------------------------------------------
# Test of:
### --- Graded tests for the 'featureBootStrap' method 
# Important Note: additional test might also be taken from our side.

# --------------------------------------------------------
folderName = '.' + os.sep + 'data' 
datasetCsvFileName = folderName + os.sep + 'spam_notSpam.csv'
# --------------------------------------------------------
print ("Testing your implementation of the 'featureBootStrap' method ...")
dataset_forTesting = loadDataset(datasetCsvFileName)
X_vectors, y_categories = separateTo_X_and_y(dataset_forTesting, 'class')
# --------------------------------------------------------

# --------------------------------------------------------
allValidVals = lambda arrValues,validVals: False not in (val in validVals for val in arrValues) 
sameIndexes = lambda obj1,obj2: False not in (obj1.index.tolist()[n] == obj2.index.tolist()[n] for n in range(len(obj1.index))) 

columns = X_vectors.columns
columns_50 = columns[:50] 
X_vectors_columns_50 = X_vectors[columns_50]
sampleRatio = 0.5
X_Sampled = featureBootStrap(X_vectors_columns_50, sampleRatio)
# --------------------------------------------------------

assert X_Sampled is not None, 'X_Sampled object not initialized'
assert isinstance(X_Sampled, pd.DataFrame), 'X_Sampled object is not a dataframe'
assert X_Sampled.empty == False, 'X_Sampled dataframe object is empty'
assert X_Sampled.shape[1] >10, 'X_Sampled dataframe object has a wrong number of sampled features'

print ("\n----> The 'featureBootStrap' test passed successfully :-) \n")
print ('\nAdditional Tests might be executed on our side')
dataset_forTesting = None
X_vectors = None
y_categories = None
columns = None
columns_50 = None
X_vectors_columns_50 = None
sampleRatio = None
X_Sampled = None
# --------------------------------------------------------

# --------------------------------------------------------
```

#### Part 3b - Building a model  - maximum possible points: 4
The following cells perform the following:
* step 1 - decision stumps (and other classification utilities) ----> run only
* step 2 - Bagging (fit)  ----> Student's implementation - total 4 points

#### Part 3b - step 1 - decision stumps (and other classification utilities) ----> run only
<img src="./images/treeStumps.jpg" alt="bagging-classifier" width="400" align='center'/>

```python
# --------------------------------------------------------
# NO IMPLEMENTAION - DO NO CHANGE
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# class 1: DecisionStump
# --------------------------------------------------------

# --------------------------------------------------------
# What does the method do?
# --- wraps an api for decision stumps, so we have a unified api
# ------------
class DecisionStump():
    def __init__(self):
        self._decisionStump = DecisionTreeClassifier(max_depth=1)
        
    def fit(self,X_train,y_train):
        self._decisionStump.fit(X_train,y_train)

    def predict(self,X_test):
        return self._decisionStump.predict(X_test)


# --------------------------------------------------------
# class2: ClassiferInstGen
# --------------------------------------------------------

# --------------------------------------------------------
# What does the method do?
# --- utility classifier to generate objects of the classification algorithm class   
# ------------
class ClassiferInstGen():
    def __init__(self,classifierPyClass):
        self._classifierPyClass = classifierPyClass
        
    def getNewClassifierPyObj(self):
        return self._classifierPyClass()

# --------------------------------------------------------
# class3: BootstrapFeatureClassifer
# --------------------------------------------------------

# --------------------------------------------------------
# What does the method do?
# --- container classifier to save model and bootstrap features    
# ------------
class BootstrapFeatureClassifer():
    def __init__(self,trainModel,featureNames):
        self.model = trainModel
        self.featureNames = featureNames

        
# --------------------------------------------------------
# method name1: trainBootstrapedClassificationModel
# --------------------------------------------------------

# --------------------------------------------------------
# What does the method do?
# --- trains a classification model which matches the sklearn API of fit and predict.
#     - It does so using the sklearn built in method
# ------------
# input parameters:
# - X_sampled_train -  a dataframe containing all feature vectors of the train set
# - y_train - a series of containing all class values per train instance
# Note: the X_sampled_train, y_train parameters could be instance bootsraped, feature bootsraped, both or none
# - classificationAlgo_pyClass - the python 'class' parameter of a classification algorithm 
#                                For instance, the above 'DecisionStump' class.
#                                Note: passing the pyton class is similar to passing a method
#                                      as a parameter. Each call to: classificationAlgo_pyClass()
#                                      creates a new object (also called instance of the class)
#                                      of the type of the class.  
# ------------
# return value:
# - featuresTrainModelObj - an object including trained model and feature names 
# --------------------- 
def trainBootstrapedClassificationModel(X_sampled_train, y_train,classificationAlgo_pyClass):
    classiferInstGen = ClassiferInstGen(classificationAlgo_pyClass)
    classificationObj = classiferInstGen.getNewClassifierPyObj()
    classificationObj.fit(X_sampled_train, y_train)
    featuresTrainModelObj = BootstrapFeatureClassifer(classificationObj,X_sampled_train.columns)
    return featuresTrainModelObj

# --------------------------------------------------------
# method name2: classifierPredict
# --------------------------------------------------------

# --------------------------------------------------------
# What does the method do?
# --- predict test examples using a classification model (which corresponds to sklearn classifier APIs)
#     - It does so using the sklearn built in method
# ------------
# input parameters:
# - X_test -  a dataframe containing feature vectors of the test set
# - classification_obj - trained classificaion model (which corresponds to sklearn classifier APIs)
# ------------
# return value:
# - yHat - a series of the predictions for each test instance  
# --------------------- trainBootstrapedFeatureModel(X_sampled_train, y_train,classificationAlgo_pyClass)
def classifierPredict(X_test, featuresTrainModelObj):
    X_adapted = X_test[featuresTrainModelObj.featureNames]
    predictions = featuresTrainModelObj.model.predict(X_adapted)
    yHat = pd.Series(data=predictions,index=X_test.index)
    return yHat

# --------------------------------------------------------
# Temporary imports:
import sklearn
from sklearn.model_selection import train_test_split
from sklearn.tree import DecisionTreeClassifier
# --------
folderName = '.' + os.sep + 'data' 
datasetCsvFileName = folderName + os.sep + 'spam_notSpam.csv'
# --------------------- 
dataset_forTesting = loadDataset(datasetCsvFileName)
X_vectors, y_categories = separateTo_X_and_y(dataset_forTesting, 'class')
y_categories = datasetCategoriesToNums(y_categories, 'spam')
X_train, X_test, y_train, y_test = trainTestSplit(X_vectors, y_categories, 0.2, 11)
X_train_1st_k = X_train.iloc[:1000,:] 
y_train_1st_k = y_train.iloc[:1000]
X_train_2nd_k = X_train.iloc[1000:2000,:] 
y_train_2nd_k = y_train.iloc[1000:2000]
classificationModel1=trainBootstrapedClassificationModel(X_train_1st_k, y_train_1st_k,DecisionStump)
classificationModel2=trainBootstrapedClassificationModel(X_train_2nd_k, y_train_2nd_k,DecisionStump)
# --------------------------------------------------------


# --------------------------------------------------------
# show usages of functions
# --------------------------------------------------------
print ('A few prediction examples:')
nExamples=5

y_hat1 = classifierPredict(X_test.iloc[:nExamples,:], classificationModel1)
y_hat2 = classifierPredict(X_test.iloc[:nExamples,:], classificationModel2)
for nInd in range(nExamples):
    print ('test instance [%d]: actual: %r, prediction(classifier1): %r, prediction(classifier2): %r' %(nInd,y_test.iloc[nInd],y_hat1.iloc[nInd],y_hat2.iloc[nInd]))
# --------------------------------------------------------
if 'datasets' in sys.modules:
    del (datasets)
if 'train_test_split' in sys.modules:
    del (train_test_split)
if 'DecisionTreeClassifier' in sys.modules:
    del (DecisionTreeClassifier)
sys_modules = list(sys.modules.keys())
for mdl in sys_modules:
    if mdl.startswith('sklearn.'):
        del(sys.modules[mdl]) 
del (sklearn)
if 'sklearn' in sys.modules:
    del (sys.modules['sklearn'])
# --------------------------------------------------------

# --------------------------------------------------------
dataset_forTesting = None
X_vectors = None
y_categories = None
X_train = None
X_test = None
y_train = None
y_test = None
y_hat1 = None
y_hat2 = None
X_train_1st_k = None 
y_train_1st_k = None
X_train_2nd_k = None 
y_train_2nd_k = None
classificationModel1=None
classificationModel2=None

# --------------------------------------------------------
```

#### Part 3b - step 2 - Bagging (fit)  ----> Student's implementation - total 4 points
<img src="./images/bags.jpg" alt="The bagging operation" width="200" align='center'/>

```python
# ------------------------------------------
# Add assistance code here IF NEEDED:



# ------------------------------------------
```

<!-- grade_id: bagging-fit, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 4
# --------------------------------------------------------

# --------------------------------------------------------
# method name: baggingFit
# --------------------------------------------------------

# --------------------------------------------------------
# The following is expected:
# --- the method needs to return the a list of models created after bootstraping instances & features.
# ------------
# input parameters:
# - X_train - a dataframe containing all feature vectors of the train set
# - y_train - a series of containing all class values per train instance
# - instanceSampleRatio - the ratio of the sampeling out of training set, 
#                        * we will pass it on as a parameter, in order to sample instances.
#                        - if instanceSampleRatio<=0, no instance bootstrap is done, and we leave
#                          the training instances with no change.
# - featureSampleRatio - the ratio of the sampeling out of feature set,
#                        * we will pass it on as a parameter, in order to sample features.
#                        - if featureSampleRatio<=0, no feature bootstrap is done, and we leave
#                          the features with no change.
# - classificationAlgo_pyClass - the python 'class' parameter of a classification algorithm 
#                                For instance, the above 'DecisionStump' class.
#                                Note: passing the pyton class is similar to passing a method
#                                      as a parameter.    
# - numModels - number of models to train in bagging
# Note: the X_train, y_train parameters could be instance bootsraped, feature bootsraped or both
# ------------
# return value:
# - a list of trained models 
#  notes:
#        * the number of models are detrminded by the input parameter 
# ---------------------
# --- General notes:
# Don't forget to remove the 'raise NotImplementedError()' line.
# In the line following the '# YOUR CODE HERE' line, add your code
# ---------------------
# --- Additional notes:
# - You must use the parameters you've got. 
# - Do NOT define the parameters in the method.
# --------------------------------------------------------
def baggingFit(X_train, y_train, instanceSampleRatio, featureSampleRatio, classificationAlgo_pyClass, numModels):
    # note that FOR EACH bootstraped train set and/or feature set you need to create a classifier using:
    # classificationModel=trainBootstrapedClassificationModel(X_train_bootstraped, y_train_corrolated_labels,classificationAlgo_pyClass)
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# --------------------------------------------------------
# - add your tests here IF NEEDED:
# Note: any additional tests are for your use only.
# Note: We depend only on tests we wrote.
# ------


# --------------------------------------------------------
```

<!-- grade_id: cell-fcf5b9299c2d8e0a, points: 4, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# Important Notes:
# - This test cell MUST NOT raise any exception
# - It must pass all tests to get the point ----- #
# - Do NOT change or delete this cell
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 4
# --------------------------------------------------------

# --------------------------------------------------------
# Test of:
### --- Graded tests for the 'baggingFit' method 
# Important Note: additional test might also be taken from our side.

# --------------------------------------------------------
# Temporary imports:
import sklearn
from sklearn.model_selection import train_test_split
from sklearn.tree import DecisionTreeClassifier
# --------
folderName = '.' + os.sep + 'data' 
datasetCsvFileName = folderName + os.sep + 'spam_notSpam.csv'
# --------------------- 
print ("Testing your implementation of the 'baggingFit' method ...")
dataset_forTesting = loadDataset(datasetCsvFileName)
X_vectors, y_categories = separateTo_X_and_y(dataset_forTesting, 'class')
y_categories = datasetCategoriesToNums(y_categories, 'spam')
X_train, X_test, y_train, y_test = trainTestSplit(X_vectors, y_categories, 0.2, 25)
# --------------------------------------------------------

# --------------------------------------------------------
allValidVals = lambda arrValues,validVals: False not in (val in validVals for val in arrValues) 
sameIndexes = lambda obj1,obj2: False not in (obj1.index.tolist()[n] == obj2.index.tolist()[n] for n in range(len(obj1.index))) 
innerElementTupple = lambda arrValues: False not in (isinstance(elem,tuple) for elem in arrValues)
# --------------------------------------------------------
instanceSampleRatio = 0.1
featureSampleRatio = 0.5
numModels = 3
classification_PyClass = DecisionStump
baggedModels =  baggingFit(X_train, y_train, instanceSampleRatio, featureSampleRatio, classification_PyClass, numModels)
# --------------------------------------------------------

print ("check basic 'baggingFit' output validation ...")
assert baggedModels is not None, 'baggedModels object not initialized'
assert isinstance(baggedModels, list), 'baggedModels object is not a list'
assert None not in baggedModels, 'baggedModels should not include None elements'
assert False not in [isinstance(elem,BootstrapFeatureClassifer) for elem in baggedModels], 'baggedModels should include only DecisionStump objects'

print ('trying to create inner bagged models prediction ...')
nExamples=20
y_hat_arr = [classifierPredict(X_test.iloc[:nExamples],model) for model in baggedModels]
assert False not in [allValidVals(y_hat.values,(0,1)) for y_hat in y_hat_arr], 'something went wrong with bagging inner models - invalid values found (valid values: 0 or 1)'

print ("----> The 'baggingFit' test passed successfully :-) \n")
print ('\nAdditional Tests might be executed on our side')

print ('\n----------------------')
print ('A few prediction examples:')
for nInd in range(nExamples):
    predictions = '; '.join('[bag-model %d]: %r' %(nPred,y_hat_arr[nPred].iloc[nInd]) for nPred in range(len(y_hat_arr)))
    print ('test instance [%d]: actual: %r, predictions: %r' %(nInd,y_test.iloc[nInd],predictions))


# --------------------------------------------------------
if 'datasets' in sys.modules:
    del (datasets)
if 'train_test_split' in sys.modules:
    del (train_test_split)
if 'DecisionTreeClassifier' in sys.modules:
    del (DecisionTreeClassifier)
sys_modules = list(sys.modules.keys())
for mdl in sys_modules:
    if mdl.startswith('sklearn.'):
        del(sys.modules[mdl]) 
del (sklearn)
if 'sklearn' in sys.modules:
    del (sys.modules['sklearn'])
# --------------------------------------------------------

# --------------------------------------------------------
dataset_forTesting = None
X_vectors = None
y_categories = None
X_train = None
X_test = None
y_train = None
y_test = None
y_hat = None

instanceSampleRatio = None
featureSampleRatio = None
numModels = None
classification_PyClass = None
baggedModels = None

nExamples = None
y_hat_arr = None
# --------------------------------------------------------

# --------------------------------------------------------
```

### Part 4 - Prediction -  maximum possible points: 3
The following cells perform the following:
* step 1 - predict (voting)    ----> Student's implementation - total 3 points

```python
# ------------------------------------------
# Add assistance code here IF NEEDED:



# ------------------------------------------
```

<!-- grade_id: bagging-predict, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 2
# --------------------------------------------------------

# --------------------------------------------------------
# method name: baggingPredict
# --------------------------------------------------------

# --------------------------------------------------------
# The following is expected:
# --- the method needs to return the predicted value for each test instances, 
#     based on the bagged trained models 
# ------------
# input parameters:
# - X_test -  a dataframe containing feature vectors of the test set
# - baggingModels - the list of bagged models returned from 'baggingFit'
# ------------
# return value:
# - y_hat - a series of the prediction, with the same index as X_test 
#  notes:
#        * the number of models are detrminded by the input parameter 
# ---------------------
# --- General notes:
# Don't forget to remove the 'raise NotImplementedError()' line.
# In the line following the '# YOUR CODE HERE' line, add your code
# ---------------------
# --- Additional notes:
# - You must use the parameters you've got. 
# - Do NOT define the parameters in the method.
# --------------------------------------------------------
def baggingPredict(X_test, baggingModels):
    # note that FOR EACH model in baggingModels, predict the value using:
    # y_test=classifierPredict(X_test, baggingModel)
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# --------------------------------------------------------
# - add your tests here IF NEEDED:
# Note: any additional tests are for your use only.
# Note: We depend only on tests we wrote.
# ------


# --------------------------------------------------------
```

<!-- grade_id: bagging-predict-test, points: 3, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL (PART 1)
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# Important Notes:
# - This test cell MUST NOT raise any exception
# - It must pass all tests to get the point ----- #
# - Do NOT change or delete this cell
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 3
# --------------------------------------------------------

# --------------------------------------------------------
# Test of:
### --- Graded tests for the 'baggingPredict' method 
# Important Note: additional test might also be taken from our side.

# --------------------------------------------------------
# Temporary imports:
import sklearn
from sklearn.model_selection import train_test_split
from sklearn.tree import DecisionTreeClassifier
# --------
folderName = '.' + os.sep + 'data' 
datasetCsvFileName = folderName + os.sep + 'spam_notSpam.csv'
# --------------------- 
dataset_forTesting = loadDataset(datasetCsvFileName)
X_vectors, y_categories = separateTo_X_and_y(dataset_forTesting, 'class')
y_categories = datasetCategoriesToNums(y_categories, 'spam')
X_train, X_test, y_train, y_test = trainTestSplit(X_vectors, y_categories, 0.2, 25)
# --------------------------------------------------------

# --------------------------------------------------------
allValidVals = lambda arrValues,validVals: False not in (val in validVals for val in arrValues) 
sameIndexes = lambda obj1,obj2: False not in (obj1.index.tolist()[n] == obj2.index.tolist()[n] for n in range(len(obj1.index))) 
innerElementTupple = lambda arrValues: False not in (isinstance(elem,tuple) for elem in arrValues)
# --------------------------------------------------------

# --------------------------------------------------------

print ("check basic 'baggingPredict' output validation (part 1) ...")
instanceSampleRatio = 0.1
featureSampleRatio = 0.5
numModels = 10
classification_PyClass = DecisionStump
baggedModels =  baggingFit(X_train, y_train, instanceSampleRatio, featureSampleRatio, classification_PyClass, numModels)
yHat = baggingPredict(X_test, baggedModels)
assert yHat is not None, 'yHat object not initialized'
assert isinstance(yHat, pd.Series), 'yHat object is not a pandas series'
assert None not in yHat.values, 'yHat should not include None elements'

print ("----> The 'baggingPredict' test passed successfully :-) \n")
print ('\nAdditional Tests might be executed on our side')
# --------------------------------------------------------
print ("check basic 'baggingPredict' output validation (part 2) ...")
assert allValidVals(yHat.values,(0,1)), 'prediction err - something went wrong with bagging inner models - invalid values found (valid values: 0 or 1)'

print ("----> The 'baggingPredict' test passed successfully :-) \n")
print ('\nAdditional Tests might be executed on our side')

print ('\n----------------------')
print ('A few prediction examples:')
nExamples = 20
for nInd in range(nExamples):
    print ('test instance [%d]: actual: %r, predictions: %r' %(nInd,y_test.iloc[nInd],yHat.iloc[nInd]))


# --------------------------------------------------------
if 'datasets' in sys.modules:
    del (datasets)
if 'train_test_split' in sys.modules:
    del (train_test_split)
if 'DecisionTreeClassifier' in sys.modules:
    del (DecisionTreeClassifier)
sys_modules = list(sys.modules.keys())
for mdl in sys_modules:
    if mdl.startswith('sklearn.'):
        del(sys.modules[mdl]) 
del (sklearn)
if 'sklearn' in sys.modules:
    del (sys.modules['sklearn'])
# --------------------------------------------------------

# --------------------------------------------------------
dataset_forTesting = None
X_vectors = None
y_categories = None
X_train = None
X_test = None
y_train = None
y_test = None
y_hat = None
yHat = None

instanceSampleRatio = None
featureSampleRatio = None
numModels = None
classification_PyClass = None
baggedModels = None

nExamples = None
y_hat_arr = None
# --------------------------------------------------------

```

### Part 5 - Evaluation - maximum possible points: 2
The following cells perform the following:
* step 1 - evaluate confusion matrix ----> run only
* step 2 - evaluate precision  ----> Student's implementation - total 1 points
* step 3 - evaluate recall     ----> Student's implementation - total 1 points

#### Part 5 - step 1 - evaluate confusion matrix ----> run only
<img src="./images/confusion_matrix.jpg" alt="confusion_matrix" width="300" align='center'/>

```python
# --------------------------------------------------------
# NO IMPLEMENTAION - DO NO CHANGE
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# methos: getConfusionMatrix
# --------------------------------------------------------

# --------------------------------------------------------
# What does the method do?
# --- compute the 4 values of the confusion matrix: TN, FP, FN, TP
#     - It does so using the sklearn built in method
# ------------
# input parameters:
# - y_hat -  a series of containing all predicted class values per test instance
# - y_test - a series of containing all class values per test instance
# ------------
# return values (comma seperated):
# - TN - True negatives - number of instances, for which the actual value (from y_test) 
#                         is 0 (negative class, not_spam in our dataset) and 
#                         the predicted value (from y_hat) is also 0.
# - FP - True negatives - number of instances, for which the actual value (from y_test) is 0, but 
#                         the predicted value (from y_hat) is 1 (positive class, spam in our dataset).
# - FN - False negatives - number of instances, for which the actual value (from y_test) is 1, but                          
#                         the predicted value (from y_hat) is 0.
# - TP - False negatives - number of instances, for which the actual value (from y_test) is 1 and                          
#                         the predicted value (from y_hat) is also 1.
# --------------------- 
def getConfusionMatrix(y_hat,y_test):
    TN = len([1 for ind in range(len(y_test)) if y_test.iloc[ind]==0 and y_hat.iloc[ind]==0])
    FP = len([1 for ind in range(len(y_test)) if y_test.iloc[ind]==0 and y_hat.iloc[ind]==1])
    FN = len([1 for ind in range(len(y_test)) if y_test.iloc[ind]==1 and y_hat.iloc[ind]==0])
    TP = len([1 for ind in range(len(y_test)) if y_test.iloc[ind]==1 and y_hat.iloc[ind]==1])
    return TN,FP,FN,TP


# --------------------------------------------------------
# Temporary imports:
import sklearn
from sklearn.model_selection import train_test_split
from sklearn.tree import DecisionTreeClassifier
# --------
folderName = '.' + os.sep + 'data' 
datasetCsvFileName = folderName + os.sep + 'spam_notSpam.csv'
# --------------------- 
dataset_forTesting = loadDataset(datasetCsvFileName)
X_vectors, y_categories = separateTo_X_and_y(dataset_forTesting, 'class')
y_categories = datasetCategoriesToNums(y_categories, 'spam')
X_train, X_test, y_train, y_test = trainTestSplit(X_vectors, y_categories, 0.2, 11)
instanceSampleRatio = 0.1
featureSampleRatio = 0.5
numModels = 10
classification_PyClass = DecisionStump
baggedModels =  baggingFit(X_train, y_train, instanceSampleRatio, featureSampleRatio, classification_PyClass, numModels)
y_hat = baggingPredict(X_test, baggedModels)
# --------------------------------------------------------

# --------------------------------------------------------
numTN,numFP,numFN,numTP = getConfusionMatrix(y_hat,y_test)
print ('confusion matrix:')
print ('\tpred-0\tpred-1')
print ('is-0\tTN=%d\tFP=%d' %(numTN,numFP))
print ('is-1\tFN=%d\tTP=%d' %(numFN,numTP))
# --------------------------------------------------------
if 'datasets' in sys.modules:
    del (datasets)
if 'train_test_split' in sys.modules:
    del (train_test_split)
if 'DecisionTreeClassifier' in sys.modules:
    del (DecisionTreeClassifier)
sys_modules = list(sys.modules.keys())
for mdl in sys_modules:
    if mdl.startswith('sklearn.'):
        del(sys.modules[mdl]) 
del (sklearn)
if 'sklearn' in sys.modules:
    del (sys.modules['sklearn'])
# --------------------------------------------------------

# --------------------------------------------------------
dataset_forTesting = None
X_vectors = None
y_categories = None
X_train = None
X_test = None
y_train = None
y_test = None

yHat = None

instanceSampleRatio = None
featureSampleRatio = None
numModels = None
classification_PyClass = None
baggedModels = None

nExamples = None

# --------------------------------------------------------
```

#### Part 5 - step 2 - evaluate precision  ----> Student's implementation - total 1 points
<img src="./images/confusion_matrix-precision-recall.jpg" alt="confusion_matrix-precision-recall" width="300" align='center'/>

```python
# ------------------------------------------
# Add assistance code here IF NEEDED:



# ------------------------------------------
```

<!-- grade_id: precision, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 5
# --------------------------------------------------------

# --------------------------------------------------------
# method name: getPrecision
# --------------------------------------------------------

# --------------------------------------------------------
# The following is expected:
# --- the method needs to return the precision for class 1 (spam in out dataset)
# ------------
# input parameters:
# - y_hat -  a series of containing all predicted class values per test instance
# - y_test - a series of containing all class values per test instance
# ------------
# return value:
# - precision value
# ---------------------
# --- General notes:
# Don't forget to remove the 'raise NotImplementedError()' line.
# In the line following the '# YOUR CODE HERE' line, add your code
# ---------------------
# --- Additional notes:
# - You must use the parameters you've got. 
# - Do NOT define the parameters in the method.
# --------------------------------------------------------

def getPrecision(y_hat,y_test):
    # note that you could use getConfusionMatrix(y_hat,y_test) to make it simpler
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# --------------------------------------------------------
# - add your tests here IF NEEDED:
# Note: any additional tests are for your use only.
# Note: We depend only on tests we wrote.
# ------


# --------------------------------------------------------
```

<!-- grade_id: precision-test, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# Important Notes:
# - This test cell MUST NOT raise any exception
# - It must pass all tests to get the point ----- #
# - Do NOT change or delete this cell
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 1
# --------------------------------------------------------

# --------------------------------------------------------
# Test of:
### --- Graded tests for the 'getPrecision' method 
# Important Note: additional test might also be taken from our side.

# --------------------------------------------------------
# Temporary imports:
import sklearn
from sklearn.model_selection import train_test_split
from sklearn.tree import DecisionTreeClassifier
# --------
folderName = '.' + os.sep + 'data' 
datasetCsvFileName = folderName + os.sep + 'spam_notSpam.csv'
# --------------------- 
dataset_forTesting = loadDataset(datasetCsvFileName)
X_vectors, y_categories = separateTo_X_and_y(dataset_forTesting, 'class')
y_categories = datasetCategoriesToNums(y_categories, 'spam')
X_train, X_test, y_train, y_test = trainTestSplit(X_vectors, y_categories, 0.2, 25)
instanceSampleRatio = 0.1
featureSampleRatio = 0.5
numModels = 5
classification_PyClass = DecisionStump
baggedModels =  baggingFit(X_train, y_train, instanceSampleRatio, featureSampleRatio, classification_PyClass, numModels)
yHat = baggingPredict(X_test, baggedModels)
# --------------------------------------------------------

# --------------------------------------------------------
allValidVals = lambda arrValues,validVals: False not in (val in validVals for val in arrValues) 
sameIndexes = lambda obj1,obj2: False not in (obj1.index.tolist()[n] == obj2.index.tolist()[n] for n in range(len(obj1.index))) 
innerElementTupple = lambda arrValues: False not in (isinstance(elem,tuple) for elem in arrValues)
# --------------------------------------------------------

# --------------------------------------------------------

print ("check basic 'getPrecision' output validation ...")
precision = getPrecision(yHat,y_test)
assert precision is not None, 'precision not initialized'
assert precision>0, 'precision should not be 0'
assert precision<=1, 'precision should not be more than 1 (=100%)'
assert precision>0.7, 'precision should not be >0.7 (more than 70%)'
print ('pricision=%r' %(precision))
print ("----> The 'getPrecision' test passed successfully :-) \n")
print ('\nAdditional Tests might be executed on our side')

print ('\n----------------------')


# --------------------------------------------------------
if 'datasets' in sys.modules:
    del (datasets)
if 'train_test_split' in sys.modules:
    del (train_test_split)
if 'DecisionTreeClassifier' in sys.modules:
    del (DecisionTreeClassifier)
sys_modules = list(sys.modules.keys())
for mdl in sys_modules:
    if mdl.startswith('sklearn.'):
        del(sys.modules[mdl]) 
del (sklearn)
if 'sklearn' in sys.modules:
    del (sys.modules['sklearn'])
# --------------------------------------------------------

# --------------------------------------------------------
dataset_forTesting = None
X_vectors = None
y_categories = None
X_train = None
X_test = None
y_train = None
y_test = None
y_hat = None
yHat = None

instanceSampleRatio = None
featureSampleRatio = None
numModels = None
classification_PyClass = None
baggedModels = None

nExamples = None
y_hat_arr = None
# --------------------------------------------------------

# --------------------------------------------------------
```

#### Part 5 - step 3 - evaluate recall     ----> Student's implementation - total 1 points
<img src="./images/confusion_matrix-precision-recall.jpg" alt="confusion_matrix-precision-recall" width="300" align='center'/>

```python
# ------------------------------------------
# Add assistance code here IF NEEDED:



# ------------------------------------------
```

<!-- grade_id: recall, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 5
# --------------------------------------------------------

# --------------------------------------------------------
# method name: getRecall
# --------------------------------------------------------

# --------------------------------------------------------
# The following is expected:
# --- the method needs to return the recall for class 1 (spam in out dataset)
# ------------
# input parameters:
# - y_hat -  a series of containing all predicted class values per test instance
# - y_test - a series of containing all class values per test instance
# ------------
# return value:
# - recall value
# ---------------------
# --- General notes:
# Don't forget to remove the 'raise NotImplementedError()' line.
# In the line following the '# YOUR CODE HERE' line, add your code
# ---------------------
# --- Additional notes:
# - You must use the parameters you've got. 
# - Do NOT define the parameters in the method.
# --------------------------------------------------------

def getRecall(y_hat,y_test):
    # note that you could use getConfusionMatrix(y_hat,y_test) to make it simpler
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# --------------------------------------------------------
# - add your tests here IF NEEDED:
# Note: any additional tests are for your use only.
# Note: We depend only on tests we wrote.
# ------


# --------------------------------------------------------
```

<!-- grade_id: recall-test, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# Important Notes:
# - This test cell MUST NOT raise any exception
# - It must pass all tests to get the point ----- #
# - Do NOT change or delete this cell
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 1
# --------------------------------------------------------

# --------------------------------------------------------
# Test of:
### --- Graded tests for the 'getRecall' method 
# Important Note: additional test might also be taken from our side.

# --------------------------------------------------------
# Temporary imports:
import sklearn
from sklearn.model_selection import train_test_split
from sklearn.tree import DecisionTreeClassifier
# --------
folderName = '.' + os.sep + 'data' 
datasetCsvFileName = folderName + os.sep + 'spam_notSpam.csv'
# --------------------- 
dataset_forTesting = loadDataset(datasetCsvFileName)
X_vectors, y_categories = separateTo_X_and_y(dataset_forTesting, 'class')
y_categories = datasetCategoriesToNums(y_categories, 'spam')
X_train, X_test, y_train, y_test = trainTestSplit(X_vectors, y_categories, 0.2, 25)
instanceSampleRatio = 0.1
featureSampleRatio = 0.5
numModels = 10
classification_PyClass = DecisionStump
baggedModels =  baggingFit(X_train, y_train, instanceSampleRatio, featureSampleRatio, classification_PyClass, numModels)
yHat = baggingPredict(X_test, baggedModels)
# --------------------------------------------------------

# --------------------------------------------------------
allValidVals = lambda arrValues,validVals: False not in (val in validVals for val in arrValues) 
sameIndexes = lambda obj1,obj2: False not in (obj1.index.tolist()[n] == obj2.index.tolist()[n] for n in range(len(obj1.index))) 
innerElementTupple = lambda arrValues: False not in (isinstance(elem,tuple) for elem in arrValues)
# --------------------------------------------------------

# --------------------------------------------------------

print ("check basic 'getRecall' output validation ...")
recall = getRecall(yHat,y_test)
assert recall is not None, 'precision not initialized'
assert recall>0, 'recall should not be 0'
assert recall<=1, 'recall should not be more than 1 (=100%)'
assert recall>0.3, 'recall should not be >0.5 (more than 50%)'
print ('recall=%r' %(recall))
print ("----> The 'getRecall' test passed successfully :-) \n")
print ('\nAdditional Tests might be executed on our side')

print ('\n----------------------')


# --------------------------------------------------------
if 'datasets' in sys.modules:
    del (datasets)
if 'train_test_split' in sys.modules:
    del (train_test_split)
if 'DecisionTreeClassifier' in sys.modules:
    del (DecisionTreeClassifier)
sys_modules = list(sys.modules.keys())
for mdl in sys_modules:
    if mdl.startswith('sklearn.'):
        del(sys.modules[mdl]) 
del (sklearn)
if 'sklearn' in sys.modules:
    del (sys.modules['sklearn'])
# --------------------------------------------------------

# --------------------------------------------------------
dataset_forTesting = None
X_vectors = None
y_categories = None
X_train = None
X_test = None
y_train = None
y_test = None
y_hat = None
yHat = None

instanceSampleRatio = None
featureSampleRatio = None
numModels = None
classification_PyClass = None
baggedModels = None

nExamples = None
y_hat_arr = None
# --------------------------------------------------------

# --------------------------------------------------------
```

### Part 6 - Bonus Question - maximum possible points: 5
The following cells perform the following:
* step 1 - evaluate accuracy ----> run only
* step 2 - split the data-set to train-set, validation-set and test-set ----> run only
* step 3 - bagging hyperparameters tuning ----> Student's implementation - total 5 points

#### Part 6 - step 1 - evaluate accuracy   ----> run only

```python
# --------------------------------------------------------
# NO IMPLEMENTAION - DO NO CHANGE
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# methos: getConfusionMatrix
# --------------------------------------------------------

# --------------------------------------------------------
# What does the method do?
# --- computes the accuracy of the classifier
# ------------
# input parameters:
# - y_hat -  a series of containing all predicted class values per test instance
# - y_test - a series of containing all class values per test instance
# ------------
# return value:
# - accuracy value
# --------------------- 
def getAccuracy(y_hat,y_test):
    correct = float(len([1 for ind in range(len(y_test)) if y_test.iloc[ind]==y_hat.iloc[ind]]))
    return correct/float(len(y_test))


# --------------------------------------------------------
# Temporary imports:
import sklearn
from sklearn.model_selection import train_test_split
from sklearn.tree import DecisionTreeClassifier
# --------
folderName = '.' + os.sep + 'data' 
datasetCsvFileName = folderName + os.sep + 'spam_notSpam.csv'
# --------------------- 
dataset_forTesting = loadDataset(datasetCsvFileName)
X_vectors, y_categories = separateTo_X_and_y(dataset_forTesting, 'class')
y_categories = datasetCategoriesToNums(y_categories, 'spam')
X_train, X_test, y_train, y_test = trainTestSplit(X_vectors, y_categories, 0.2, 11)
instanceSampleRatio = 0.1
featureSampleRatio = 0.5
numModels = 5
classification_PyClass = DecisionStump
baggedModels =  baggingFit(X_train, y_train, instanceSampleRatio, featureSampleRatio, classification_PyClass, numModels)
y_hat = baggingPredict(X_test, baggedModels)
# --------------------------------------------------------

# --------------------------------------------------------
accuracy = getAccuracy(y_hat,y_test)
print ('accuracy: %r' %(accuracy))
# --------------------------------------------------------
if 'datasets' in sys.modules:
    del (datasets)
if 'train_test_split' in sys.modules:
    del (train_test_split)
if 'DecisionTreeClassifier' in sys.modules:
    del (DecisionTreeClassifier)
sys_modules = list(sys.modules.keys())
for mdl in sys_modules:
    if mdl.startswith('sklearn.'):
        del(sys.modules[mdl]) 
del (sklearn)
if 'sklearn' in sys.modules:
    del (sys.modules['sklearn'])
# --------------------------------------------------------

# --------------------------------------------------------
dataset_forTesting = None
X_vectors = None
y_categories = None
X_train = None
X_test = None
y_train = None
y_test = None

yHat = None

instanceSampleRatio = None
featureSampleRatio = None
numModels = None
classification_PyClass = None
baggedModels = None

nExamples = None

# --------------------------------------------------------
```

#### Part 6 - step 2 - split the data-set to train-set, validation-set and test-set ----> run only
<img src="./images/train-validation-test-split.png" alt="train-validation-test split" width="300" align='center'/>

```python
# --------------------------------------------------------
# NO IMPLEMENTAION - DO NO CHANGE
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# method name1: trainValidationTestSplit
# --------------------------------------------------------

# --------------------------------------------------------
# What does the method do?
# --- the method split the input dataset into train and test. 
#     - It does so using the sklearn built in method
# ------------
# input parameters:
# - X_vectors - a dataframe containing all feature vectors. 
# - y_categories - a series of containing all class values per instance.
# - test_size_ratio - a number (0<number<1) of the wanted ratio of the dataset out of the dataset 
# - rand_state - a number, in order to guarantee reproducible results 
# ------------
# return values (comma separated):
# - X_train -  a dataframe containing all feature vectors of the train set
# - X_test -  a dataframe containing all feature vectors of the test set
# - y_train - a series of containing all class values per train instance
# - y_test - a series of containing all class values per test instance
# --------------------- 
def trainValidationTestSplit(X_vectors, y_categories, validation_or_test_size_ratio, rand_state):
    X_train, X_test, y_train, y_test = train_test_split(X_vectors, y_categories, test_size=validation_or_test_size_ratio, random_state=rand_state)
    validationRatio = validation_or_test_size_ratio / (1-validation_or_test_size_ratio)
    X_train, X_validation, y_train, y_validation =  train_test_split(X_train, y_train, test_size=validationRatio, random_state=rand_state)
    return X_train, X_validation, X_test, y_train, y_validation, y_test  

# --------------------------------------------------------
# Temporary imports:
import sklearn
from sklearn.model_selection import train_test_split
# --------
folderName = '.' + os.sep + 'data' 
datasetCsvFileName = folderName + os.sep + 'spam_notSpam.csv'
# --------------------- 
dataset_forTesting = loadDataset(datasetCsvFileName)
X_vectors, y_categories = separateTo_X_and_y(dataset_forTesting, 'class')
y_categories = datasetCategoriesToNums(y_categories, 'spam')
X_train, X_validation, X_test, y_train, y_validation, y_test =  trainValidationTestSplit(X_vectors, y_categories, 0.2, 43)

# --------------------------------------------------------

# --------------------------------------------------------
if 'datasets' in sys.modules:
    del (datasets)
if 'train_test_split' in sys.modules:
    del (train_test_split)
sys_modules = list(sys.modules.keys())
for mdl in sys_modules:
    if mdl.startswith('sklearn.'):
        del(sys.modules[mdl]) 
del (sklearn)
if 'sklearn' in sys.modules:
    del (sys.modules['sklearn'])
# --------------------------------------------------------

# --------------------------------------------------------
# display train-test split information
# --------------------------------------------------------
print ('Information after train-test split:')
print('* The train-set includes %d instances and %d corresponding categories' %(X_train.shape[0],y_train.shape[0]))
print('* The validation-set includes %d instances and %d corresponding categories' %(X_validation.shape[0],y_validation.shape[0]))
print('* The test-set includes %d instances and %d corresponding categories' %(X_test.shape[0],y_test.shape[0]))

```

#### Part 6 - step 3 - bagging hyperparameters tuning (using grid search) ----> Student's implementation - total 5 points
<img src="./images/grid_search.png" alt="train-validation-test split" width="100" align='center'/>

```python
# ------------------------------------------
# Add assistance code here IF NEEDED:



# ------------------------------------------
```

<!-- grade_id: hyperparametes-tune, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 5
# --------------------------------------------------------

# --------------------------------------------------------
# method name: baggingFit
# --------------------------------------------------------

# --------------------------------------------------------
# The following is expected:
# --- the method chooses the best permutation of bagging model hyperparameters, using grid search
# ------------
# input parameters:
# - X_train - a dataframe containing all feature vectors of the train set
# - y_train - a series of containing all class values per train instance
# - instanceSampleRatio - the ratio of the sampeling out of training set, 
#                        * we will pass it on as a parameter, in order to sample instances.
#                        - if instanceSampleRatio<=0, no instance bootstrap is done, and we leave
#                          the training instances with no change.
# - featureSampleRatio - the ratio of the sampeling out of feature set,
#                        * we will pass it on as a parameter, in order to sample features.
#                        - if featureSampleRatio<=0, no feature bootstrap is done, and we leave
#                          the features with no change.
# - classificationAlgo_pyClass_Arr - an array of the python 'class' parameters of a classification algorithm to choose from,
#                                For instance, the above [DecisionStump, NaiveBayes].
#                                Note: passing the pyton class is similar to passing a method
#                                      as a parameter.    
# - numModels_Arr - an array of the options of the number of bagging models to train, for which we need to choose from
#                   For instance: [3,5,10]
# Note: the X_train, y_train parameters could be instance bootsraped, feature bootsraped or both
# ------------
# return values (comma seperated):
# - allBaggedModels - an array of all trainedBaggedModels,
# - best_baggedModels - array of the baggedModels ensembles (from baggeningFit), reaching the highest accuracy
# - bestAccuracy - the highest accuracy (number) of all baggedModels ensemble
# - best_classificationAlgo_pyClass - the classificationAlgo_pyClass (e.g. DecisionStump) of
#                                     baggedModels ensemble with highest accuracy
# - best_numModels - the num of 'bagged' Models (e.g. 5) of
#                    baggedModels ensemble with highest accuracy
# ---------------------
# --- General notes:
# Don't forget to remove the 'raise NotImplementedError()' line.
# In the line following the '# YOUR CODE HERE' line, add your code
# ---------------------
# --- Additional notes:
# - You must use the parameters you've got. 
# - Do NOT define the parameters in the method.
# --------------------------------------------------------
def gridSearchBaggingModel(X_train, y_train, X_validation, y_validation, instanceSampleRatio, featureSampleRatio, classificationAlgo_pyClass_Arr, numModels_Arr):
    # hint: for each permutation, use assistance functions:
    #       - baggingFit, baggingPredict, getAccuracy 
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# --------------------------------------------------------
# - add your tests here IF NEEDED:
# Note: any additional tests are for your use only.
# Note: We depend only on tests we wrote.
# ------


# --------------------------------------------------------
```

<!-- grade_id: hyperparametes-tune-test, points: 5, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# Important Notes:
# - This test cell MUST NOT raise any exception
# - It must pass all tests to get the point ----- #
# - Do NOT change or delete this cell
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 5
# --------------------------------------------------------

# --------------------------------------------------------
# Test of:
### --- Graded tests for the 'getRecall' method 
# Important Note: additional test might also be taken from our side.

# --------------------------------------------------------
# Temporary imports:
import sklearn
from sklearn.model_selection import train_test_split
from sklearn.tree import DecisionTreeClassifier
from sklearn.naive_bayes import GaussianNB
# --------
folderName = '.' + os.sep + 'data' 
datasetCsvFileName = folderName + os.sep + 'spam_notSpam.csv'
# --------------------- 
dataset_forTesting = loadDataset(datasetCsvFileName)
X_vectors, y_categories = separateTo_X_and_y(dataset_forTesting, 'class')
y_categories = datasetCategoriesToNums(y_categories, 'spam')
X_train, X_validation, X_test, y_train, y_validation, y_test =  trainValidationTestSplit(X_vectors, y_categories, 0.2, 19)
# --------------------------------------------------------

# --------------------------------------------------------
allValidVals = lambda arrValues,validVals: False not in (val in validVals for val in arrValues) 
sameIndexes = lambda obj1,obj2: False not in (obj1.index.tolist()[n] == obj2.index.tolist()[n] for n in range(len(obj1.index))) 
innerElementTupple = lambda arrValues: False not in (isinstance(elem,tuple) for elem in arrValues)
# --------------------------------------------------------

# --------------------------------------------------------
instanceSampleRatio = 0.1
featureSampleRatio = 0.5
numModels_Arr = [5,10,40]
classificationAlgo_pyClass_Arr = [DecisionStump,GaussianNB]
allBaggedModels,best_baggedModels,bestAccuracy, best_classificationAlgo_pyClass,best_numOfModels = gridSearchBaggingModel(
                        X_train, y_train, X_validation, y_validation, instanceSampleRatio, featureSampleRatio, 
                                                         classificationAlgo_pyClass_Arr, numModels_Arr)

print ("check basic 'gridSearchBaggingModel' output validation ...")
assert allBaggedModels is not None, 'allBaggedModels not initialized'
assert best_baggedModels is not None, 'best_baggedModels not initialized'
assert bestAccuracy is not None, 'bestAccuracy not initialized'
assert best_classificationAlgo_pyClass is not None, 'best_classificationAlgo_pyClass not initialized'
assert best_numOfModels is not None, 'best_numOfModels not initialized'
assert isinstance(allBaggedModels, list), 'allBaggedModels object is not a list'
assert isinstance(best_baggedModels, list), 'best_baggedModels object is not a list'
assert bestAccuracy>0, 'bestAccuracy should not be 0'
assert bestAccuracy<=1, 'bestAccuracy should not be more than 1 (=100%)'
assert bestAccuracy>0.6, 'bestAccuracy should not be >0.6 (more than 60%)'
assert None not in allBaggedModels, 'allBaggedModels should not include None elements'
assert None not in best_baggedModels, 'best_baggedModels should not include None elements'

print ("----> The 'gridSearchBaggingModel' test passed successfully :-) \n")
print ('\nAdditional Tests might be executed on our side')

print ('\n----------------------')
print ('Best bagged Classification algo: %r, num of models: %r, accuracy: %r' %(best_classificationAlgo_pyClass.__name__,best_numOfModels,bestAccuracy))


# --------------------------------------------------------
if 'datasets' in sys.modules:
    del (datasets)
if 'train_test_split' in sys.modules:
    del (train_test_split)
if 'DecisionTreeClassifier' in sys.modules:
    del (DecisionTreeClassifier)
sys_modules = list(sys.modules.keys())
for mdl in sys_modules:
    if mdl.startswith('sklearn.'):
        del(sys.modules[mdl]) 
del (sklearn)
if 'sklearn' in sys.modules:
    del (sys.modules['sklearn'])
# --------------------------------------------------------

# --------------------------------------------------------
dataset_forTesting = None
X_vectors = None
y_categories = None
X_train = None
X_test = None
y_train = None
y_test = None
y_hat = None
yHat = None

instanceSampleRatio = None
featureSampleRatio = None
numModels = None
classification_PyClass = None
baggedModels = None

nExamples = None
y_hat_arr = None
# --------------------------------------------------------

```